# Quantum Open Source Foundation Monthly Challenge
## Quantum Computing with Qiskit

Welcome to the qosf.org monthly challenge! This notebook will guide you through installing Qiskit and performing a simple quantum computing task: applying a controlled-X (CX) gate from qubit 0 to qubit 1 and reading the output.

### Step 1: Install Qiskit
First, we need to install Qiskit, a comprehensive open-source quantum computing framework.
Run the following cell to install Qiskit.

```python
# If you run this notebook on mac, you need to use pip3 instead
!pip install "qiskit[visualization]"
!pip install matplotlib
!pip install pylatexenc
```

### Step 2: Import Qiskit and Create a Quantum Circuit
Next, let's import Qiskit and create a quantum circuit with 2 qubits and 2 classical bits.

```python
from qiskit import QuantumCircuit
from qiskit.visualization import plot_histogram

# Create a Quantum Circuit with 2 qubits and 2 classical bits
qc = QuantumCircuit(2, 2)
```

### Step 3: Apply a Controlled-X (CX) Gate
Now, apply a controlled-X (CX) gate from qubit 0 to qubit 1.

<!-- grade_id: answer1, points: 0, grade: false, solution: true -->
```python
# Write your code here and don't change the variable or circuit name (qc)

# Apply a controlled-X (CX) gate from qubit 0 to qubit 1 and then measure the quibits

# YOUR CODE HERE
qc.cx(0, 1)
qc.measure([0, 1], [0, 1])

# Now let's draw the circuit
qc.draw(output='text')
```

Output:
```
┌─┐   
q_0: ──■──┤M├───
     ┌─┴─┐└╥┘┌─┐
q_1: ┤ X ├─╫─┤M├
     └───┘ ║ └╥┘
c: 2/══════╩══╩═
           0  1
```

### Step 4: Check Your Results
To verify your implementation, run the following cell to check your circuit

<!-- grade_id: test1, points: 1, grade: true, solution: false -->
```python

assert str(qc.data[0][0]) == "Instruction(name='cx', num_qubits=2, num_clbits=0, params=[])"
print("Done!")
```

Output:
```
Done!
```

Congratulations! You've completed the challenge. Feel free to explore more with Qiskit and quantum computing.